## CAS Data Science - Standortbestimmung in Python
#### FHNW - HS2019

Einige Beispiele mit dem Umgang von Python, Pandas und MatPlotLib werden aufgezeigt.
Dieses Warm-Up sollte die Konzepte revidieren, die in den empfohlenen Dokumente im Voraus geschickt worden sind. 

```python
#Importieren von Libraries
import numpy as np
import pandas as pd
import matplotlib.mlab as mlab
import matplotlib.pyplot as plt
import time
from nose.tools import assert_equal
```

#### Basics: Print-Funktion

```python
le_string = 'p,k,f,n,f,n,f,c,n,w,e,?,k,y,w,n,p,w,o,e,w,v,d'
zahlen_als_liste = [1,2,3,4,5]
sieben = 7
```

<!-- grade_id: 4, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

#### Anwendung und Definition von Funktionen

Führen Sie mittels einer for-Schleife die Fibonacci-Folge fort. Nutzen Sie dazu die range()-Funktion in der Schleifen-Deklaration und die append-Funktion um die Liste zu ergänzen. def definiert eine Funktion, welche Sie weiter unten ausführen können.

<!-- grade_id: 6, points: 0, grade: false, solution: true -->
```python
def fibonacci(n):
    sequence = [0,1]
    ### BEGIN SOLUTION
   
    ### END SOLUTION
    return sequence
```

```python
#Ausgabe ihrer Fibonacci-Funktion
print(fibonacci(5))
print(fibonacci(10))
```

In der nachfolgenden Zelle wird ihre Fibonacci-Funktion überprüft. Diese Tests werden auch beim submitten ihrer Aufgabe durchgeführt. Wenn die Tests fehlschlagen, dann ist ihre Lösung nicht korrekt.

<!-- grade_id: 8, points: 1, grade: true, solution: false -->
```python
assert_equal(fibonacci(5),[0, 1, 1, 2, 3])
assert_equal(fibonacci(10),[0, 1, 1, 2, 3, 5, 8, 13, 21, 34])
```

#### Listen

Listen müssen nicht den gleichen Datentyp beinhalten. Sie können auf Listen ähnliche Zugriffe wie auf NumPy Arrays durchführen.

```python
mixed_list = ['a', 1, 2.3, True, [1, 'b']]
print(mixed_list)
```

<!-- grade_id: 10, points: 0, grade: false, solution: true -->
```python
#Nutzen Sie die Slice-Notation, seq[start:stop], um die Elemente mit den Indizes 3-5 (resultat1), 
#sowie jedes zweite (die ungraden) Element (resultat2) des Objektes die_liste zu erhalten.
#Geben Sie diese beiden Resultate als Tupel (resultat1, resultat2) zurück.
def show_slice():
    die_liste = le_string.split(',')
    ### BEGIN SOLUTION
    
    ### END SOLUTION
    return (resultat1, resultat2)   
```

```python
print(show_slice())
```

<!-- grade_id: 12, points: 1, grade: true, solution: false -->
```python
# Hier prüfen wir wieder ihre Implementation der show_slice Funktion
res1, res2 = show_slice()
assert_equal(res1,['n', 'f'])
assert_equal(res2, ['k', 'n', 'n', 'c', 'w', '?', 'y', 'n', 'w', 'e', 'v'])
```

#### Dictionaries

Dictionaries sind auch bekannt als  "mappings" oder "associative arrays." Erstellt werden sie mit {}-Klammern.

```python
# Erstellen eines Dictionary
ages = {"Rick": 46, "Bob": 86, "Fred": 21}
for name in ages:
    print(name," ist ",ages[name]," alt.")
```

<!-- grade_id: 14, points: 1, grade: true, solution: true -->
```python
#Benutzen Sie die len() Funktion um die Länge des Dictionaries 'ages' und des Strings 'le_string' auszugeben.
### BEGIN SOLUTION

### END SOLUTION
```

#### Plotting-Funktionen

Nun schauen wir uns ein einfaches Beispiel eines Plots an.

```python
#Mittels matplotlib können Sie sehr einfach Plots generieren.
#Der folgende Plot visualisiert den Anstieg der Fibonacci-Folge gegenüber n!, der Fakultät.

x = np.arange(10)
facts = [1]
for i in x[0:-1]+1:
    facts.append( i * facts[i-1])
    
fibs = fibonacci(10)

print(x)
print(facts)
print(fibs)

plt.plot(x,facts,label="factorial")
plt.plot(x,fibs,label="Fibonacci")
plt.xlabel("n")
plt.legend()
```

Plotten Sie dasselbe Resultat, aber mit einer logarithmischen Skala.

<!-- grade_id: 17, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

#### Numpy

Numpy-Arrays haben bereits viele Funktionen vordefiniert und sind mächtige Werkzeuge im Umgang mit Matrizen.
Sie sind extrem Speicheroptimiert und arbeiten vektor-basiert, wie Sie hier sehen:

```python
# NumPy Beispiele
A = np.ones((3,3))
B = np.array([[2,2,2],[2,2,2],[2,2,2]])
print(A)
print(B)
print(A+B)
print(A*B)          #Skalar und keine Matrix-Multiplikation
print(np.dot(A, B)) # Matrix-Multiplikation
print(27*A)
```

Berechnen Sie element-weise die Matrix P^Q, also so dass P[0]^Q[0], P[1]^Q[1], ...usw, nur mit den Funktionalitäten von NumPy-Arrays, also ohne einen Loop zu schreiben.

<!-- grade_id: 25, points: 1, grade: true, solution: true -->
```python
P = np.arange(6).reshape(2,3)
Q = np.arange(6).reshape(2,3)
### BEGIN SOLUTION

### END SOLUTION
print(P)
```

Filtern Sie alle Elemente aus dem Array big, welche grösser als 2 sind, ohne eine Schleife zu verwenden.<br>Hint: Boolean array indexing

```python
big = np.linspace(1, 10) #Gleichmässige Verteilung zwischen 1 und 10
print(big)
```

<!-- grade_id: 22, points: 1, grade: true, solution: true -->
```python
bool_idx = (big > 2)
### BEGIN SOLUTION

### END SOLUTION
```

#### Pandas

Nachfolgend wird die Histogramm Funktion mittels eines Random Arrays gezeigt: np.random.normal() - Der Input dieser Funktion ist mean, std und die Anzahl Elemente

Mittels Pandas können auch numpy arrays manipuliert werden. Fügen sie in einem Pandas Dataframe einen numpy array ein, der gleichverteilte Daten enthält.


```python
#ANFANG LÖSUNG

#ENDE LÖSUNG

# mit Pandas kann man leicht die Verteilung kontrollieren: 
myDf.describe()
```

Zeigen Sie jetzt ein Histogramm der Werte in Ihr Array: 

```python
#ANFANG LÖSUNG

#ENDE LÖSUNG
```

Wenn Sie weitere Tutorials und Informationen suchen, dann werden Sie im Internet fündig. Einige gute Beispiele: <br>
http://nbviewer.jupyter.org/github/gumption/Python_for_Data_Science/blob/master/Python_for_Data_Science_all.ipynb <br>
http://nbviewer.jupyter.org/gist/rpmuller/5920182

```python

```